---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise to test the potential of DC.

`@instructions`
print a statement called "Hello, World!"

`@hint`
use the print() statement and pass a string

`@pre_exercise_code`
```{python}
print("...statement...")
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("Hello, World!")
```

`@sct`
```{python}

```
